# Homework 2


## Problem 1 (Dictionary and String Methods):

You have a list of strings which contain student record. Each string contains the student's name, and their grades in Maths, Physics and Chemistry. The grades can be **floating values**. These values (names and grades) are separated by comma. You are required to save the record in a dictionary data type. The user then enters a student's name. Output the average percentage marks obtained by that student, correct to two decimal places.

Example:

    input:
        grades = ["Jake, 99, 70, 50", "Dennis, 100, 100, 98"]
        name = "Dennis"
    output: "The average grade of Dennis is 99.33"

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(grades, name):
    output = ""
    
    grade_D = {}
    for student in grades:
        parsed_grades = student.strip("\n").split(",")
        names = parsed_grades[0]

        grade1 = float(parsed_grades[1])
        grade2 = float(parsed_grades[2])
        grade3 = float(parsed_grades[3])
        avg_grade = (grade1+grade2+grade3)/3

        if names in grade_D:
            grade_D[names].append(avg_grade)
        else:
            grade_D[names] = [avg_grade]
    value = grade_D.get(name)
    value_float = []
    for item in value:
        value_float.append(item)

    output = "The average grade of %s is %0.2f" %(name, value_float[0])

    return(output)
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 15, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 54", "Dennis, 100, 100, 98"]
name = "Dennis"
assert problem2(grades, "Dennis") == "The average grade of Dennis is 99.33"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 15, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 55", "Dennis, 100, 100, 98"]
name = "Jake"
assert problem2(grades, "Jake") == "The average grade of Jake is 74.67"
```

## Problem 2 (List Methods):

You are given N positive numbers. Store them in a list and find the second largest number.
(If there are two largest numbers, the second largest number must be smaller than both of them)

Example:

    input: int_list = [1,3,1,4,4,2]
    ouput: 3

<!-- grade_id: cell-90ba74290831cf3b, points: 0, grade: false, solution: true -->
```python
#Write your code here:
def problem3(int_list):
    second_largest = 0

    int_list.sort()
    for i in int_list:
        if int_list[-i-1] < int_list[-i]:
            second_largest = int_list[-i-1]
            break
    return(second_largest)
```

<!-- grade_id: cell-f14a368424f5d8fc, points: 15, grade: true, solution: false -->
```python
assert problem3([1,3,1,4,4,2]) == 3
```

<!-- grade_id: cell-b28263287d112648, points: 15, grade: true, solution: false -->
```python
assert problem3([1,2,3]) == 2
```

## Problem 3 (All-in-one: list + string + file read):

We downloaded several stock's daily level data for you in ./data/hw2_data/stock_data/. You will write code that takes stock tickers as inputs and output one dictinoary. The dictionary uses the stock tickers as keys and the values are these stocks' average open prices in 2017 (until the end of the file).

Exmaple:

    input: tickers = ["AAPL"]
    output: {'AAPL': 143.465799260355}


<!-- grade_id: cell-25ac6e4633eae166, points: 0, grade: false, solution: true -->
```python
def problem4(tickers):
    file_location = "./data/stock_data/"
    return_dict = {}
    
    for ticker in tickers:
        f = open("./data/stock_data/%s.csv" %ticker, "r")
        lines = f.readlines()
        symbol = []
        for i in lines[1:]:
            parsed_symbol = i.strip("\n").split(",")
            symbol.append(float(parsed_symbol[2]))
        avg_open_price = sum(symbol)/len(symbol)
        return_dict[ticker] = avg_open_price
    
    return(return_dict)
```

<!-- grade_id: cell-f824aae0b179bccb, points: 20, grade: true, solution: false -->
```python
assert problem4(["AAPL"]) == {'AAPL': 143.465799260355}
```

<!-- grade_id: cell-49c8a58cfdfc3878, points: 20, grade: true, solution: false -->
```python
assert problem4(["AAPL", "GOOG"]) == {'AAPL': 143.465799260355, 'GOOG': 884.3826611893493}
```

```python

```